# Ejemplo: Assignment con Tests Ocultos

Este notebook muestra cómo configurar correctamente las celdas para que los tests NO aparezcan en el feedback de los estudiantes.

## Ejercicio 1: Concatenación de Arrays

Crea dos arrays verticales de NumPy y realiza concatenaciones horizontal y vertical.

**Variables a crear:**
- `array1`: Array vertical [[1], [2], [3]]
- `array2`: Array vertical [[4], [5], [6]]
- `concatenacion_horizontal`: Concatenación horizontal de array1 y array2
- `concatenacion_vertical`: Concatenación vertical de array1 y array2

<!-- grade_id: ejercicio1_respuesta, points: 0, grade: false, solution: true -->
```python
import numpy as np

# ESCRIBE TU CÓDIGO AQUÍ
### BEGIN SOLUTION
array1 = np.array([[1], [2], [3]])
array2 = np.array([[4], [5], [6]])
concatenacion_horizontal = np.concatenate((array1, array2), axis=1)
concatenacion_vertical = np.vstack((array1, array2))
### END SOLUTION
```

<!-- grade_id: ejercicio1_test_publico, points: 1, grade: true, solution: false -->
```python
# Test público (los estudiantes VEN este código en el feedback)
print("✓ Test público: Verificando que las variables existan...")
assert 'array1' in dir(), "Debes crear la variable 'array1'"
assert 'array2' in dir(), "Debes crear la variable 'array2'"
assert 'concatenacion_horizontal' in dir(), "Debes crear la variable 'concatenacion_horizontal'"
assert 'concatenacion_vertical' in dir(), "Debes crear la variable 'concatenacion_vertical'"
print("✓ Todas las variables existen")
```

<!-- grade_id: ejercicio1_test_oculto, points: 1.5, grade: true, solution: false -->
```python
# Tests visibles (los estudiantes VEN solo el resultado pass/fail, NO el código)
assert isinstance(array1, np.ndarray), "array1 debe ser un array de NumPy"
assert isinstance(array2, np.ndarray), "array2 debe ser un array de NumPy"
assert array1.shape == array2.shape, "Los arrays deben tener la misma forma"

### BEGIN HIDDEN TESTS
# Tests completamente ocultos (los estudiantes NO VEN este código en el feedback)
# Solo ven: "✓ Test pasado" o "✗ Test fallado" con el mensaje de error

# Forma vertical: debe ser n×1
assert array1.shape[1] == 1, "Los arrays deben ser verticales (n×1)"

# Verifica la concatenación horizontal
esperado_horizontal = np.hstack((array1, array2))
assert np.array_equal(concatenacion_horizontal, esperado_horizontal), \
    "La concatenación horizontal no es correcta"

# Verifica la concatenación vertical
esperado_vertical = np.vstack((array1, array2))
assert np.array_equal(concatenacion_vertical, esperado_vertical), \
    "La concatenación vertical no es correcta"
### END HIDDEN TESTS
```

---

## ⚙️ CONFIGURACIÓN IMPORTANTE

### Para OCULTAR el código de tests del feedback:

La celda de test DEBE tener este metadata en Jupyter:

```json
{
  "nbgrader": {
    "grade": true,
    "grade_id": "nombre_unico",
    "locked": true,          ← ESTO ES CRÍTICO
    "points": 1.5,
    "schema_version": 3,
    "solution": false,
    "task": false
  }
}
```

### Pasos para configurar en Jupyter Notebook:

1. **Instalar la extensión de nbgrader:**
   ```bash
   jupyter nbextension install --sys-prefix --py nbgrader --overwrite
   jupyter nbextension enable --sys-prefix --py nbgrader
   jupyter serverextension enable --sys-prefix --py nbgrader
   ```

2. **Abrir el notebook en Jupyter**

3. **En la toolbar verás: `View > Cell Toolbar > Create Assignment`**

4. **Para cada celda de test:**
   - Selecciona la celda
   - En el dropdown de nbgrader, selecciona: **"Autograded tests"**
   - Configura los puntos
   - **MARCA el checkbox "Lock cell"** ← ESTO OCULTA EL CÓDIGO

5. **Usa los delimitadores para tests completamente ocultos:**
   ```python
   # Tests que los estudiantes ven
   assert x > 0
   
   ### BEGIN HIDDEN TESTS
   # Estos tests NO aparecen en el feedback
   # Solo muestran pass/fail
   assert x == 42
   ### END HIDDEN TESTS
   ```

### ¿Qué ven los estudiantes en el feedback?

**CON `"locked": true` y `BEGIN HIDDEN TESTS`:**
- ✓ Test pasado / ✗ Test fallado
- Mensaje de error (si falló)
- Puntaje obtenido
- **NO** ven el código de los asserts

**SIN `"locked": true`:**
- ven TODO el código incluyendo los asserts
- ven exactamente qué valores se esperan
- pueden copiar la solución de los tests

### Verificación del metadata

Ejecuta esta celda para verificar que el metadata está correcto:

```python
import json
import nbformat

# Cargar este notebook
notebook_path = 'ejemplo_assignment_con_tests_ocultos.ipynb'

try:
    with open(notebook_path, 'r', encoding='utf-8') as f:
        nb = nbformat.read(f, as_version=4)
    
    print("🔍 VERIFICACIÓN DE METADATA DE CELDAS DE TEST")
    print("="*70)
    
    test_cells_found = 0
    locked_tests = 0
    unlocked_tests = 0
    
    for i, cell in enumerate(nb.cells):
        if cell.cell_type == 'code':
            metadata = cell.get('metadata', {})
            nbgrader = metadata.get('nbgrader', {})
            
            if nbgrader.get('grade', False):
                test_cells_found += 1
                cell_id = cell.get('id', f'cell-{i}')
                grade_id = nbgrader.get('grade_id', 'sin-id')
                locked = nbgrader.get('locked', False)
                points = nbgrader.get('points', 0)
                
                status = "✅ LOCKED (tests ocultos)" if locked else "❌ UNLOCKED (tests visibles)"
                
                if locked:
                    locked_tests += 1
                else:
                    unlocked_tests += 1
                
                print(f"\nCelda {i} (id: {cell_id})")
                print(f"  grade_id: {grade_id}")
                print(f"  points: {points}")
                print(f"  locked: {locked} {status}")
                
                # Verificar si usa BEGIN HIDDEN TESTS
                if '### BEGIN HIDDEN TESTS' in cell.source:
                    if '### END HIDDEN TESTS' in cell.source:
                        print(f"  ✓ Usa delimitadores BEGIN/END HIDDEN TESTS correctamente")
                    else:
                        print(f"  ⚠️  Falta ### END HIDDEN TESTS")
    
    print("\n" + "="*70)
    print(f"📊 RESUMEN:")
    print(f"  Total de celdas de test: {test_cells_found}")
    print(f"  Tests ocultos (locked=true): {locked_tests}")
    print(f"  Tests visibles (locked=false): {unlocked_tests}")
    print("="*70)
    
    if unlocked_tests > 0:
        print("\n⚠️  ADVERTENCIA: Hay celdas de test sin locked=true")
        print("   Los estudiantes verán el código completo de los tests")
        print("   Para ocultarlos: View > Cell Toolbar > Create Assignment")
        print("   Luego marca el checkbox 'Lock cell' en cada celda de test")
    else:
        print("\n✅ ¡Perfecto! Todas las celdas de test están configuradas correctamente")

except FileNotFoundError:
    print(f"❌ No se encontró el archivo: {notebook_path}")
except Exception as e:
    print(f"❌ Error: {e}")
```